The targeted cytosolic degradation of class I histone deacetylases is essential for efficient alphaherpesvirus replication

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

In this study, the authors propose that HSV-1 infection degrades the class I histone deacetylases HDAC1 and HDAC2. The MDM2 E3 ubiquitin ligase from the DNA damage response pathway is responsible for ubiquitinating these HDACs that are subsequently degraded via proteasomes. The authors hypothesize that HDAC degradation will cause hyperacetylation of viral chromatin and enable viral gene transcription.

Strengths:

The ubiquitination of HDAC1 & HDAC2 by Mdm2 and the mapping studies are clear.

Comments on revised version:

The authors enhanced their manuscript by more supportive data and providing clarification and the necessary corrections. However, a few more issues pertain:

(1) In Figure 4j at 2 h post-infection we typically see the input virus and not progeny virus production. The input seems to have about 1-log difference that is expected to impact the results.

We sincerely appreciate the reviewer's valuable comments regarding the timing notation. It should be noted that the "2 h" indicated in Figure 4j does not refer to two hours after the start of viral infection, but rather to two hours following medium replacement—after the virus has completed adsorption and internalization at 37°C (typically taking 2 hours), with this moment defined as the new time zero point (t = 0 h). Thus, this corresponds to approximately 4 hours post-infection (4 hpi). All subsequent sampling time points (4, 6, 12, and 24 h) are consistently defined according to this same system. This temporal framework aligns with previous studies: Nobe et al. (mBio 2025; DOI: 10.1128/mbio.00280-25) have clearly demonstrated that newly generated viral particles can be detected as early as 4 hours after HSV-1 infection, supporting the possibility of early progeny virus production at this time point in our experiment. We have accordingly revised the figure legend for Figure 4j to explicitly state the time reference ("t = 0 h defined as time of medium replacement post-adsorption") and added detailed procedural descriptions in the Methods section regarding adsorption, medium change, and sample collection time points to ensure clarity and reproducibility of the timing protocol.

(2) Figs 1A, 1E, 2H it seems unclear why ICP4 becomes detectable at 12 h post-infection in HeLa cells? How about other a-genes? How about other cells? ICP4 is typically detectable within 2-3 h post-infection.

We sincerely appreciate the valuable comments provided by the reviewers. Regarding the observation that ICP4 was detected only after 12 hours post-infection in HeLa cells, we re-evaluated our experimental conditions and reviewed relevant literature. The results indicate that at a higher multiplicity of infection (MOI = 5), ICP4 can indeed be reliably detected in HeLa cells as early as 2 hours post-infection (Author response image 1). Notably, Fouad S. El-Mayet et al. reported that under MOI = 1, ICP4 could not be detected until 8 hours after HSV-1 infection of mouse neuroblastoma Neuro-2A cells Figure 5A (Fouad S. El-Mayet et al., Antiviral Research, 2024, DOI: 10.1016/j.antiviral.2024.105870), although their early protein VP16 showed positive expression as early as 4 hours post-infection. This time difference is closely related to cell type: Neuro-2A is a highly susceptible neuronal cell line for HSV-1, exhibiting significantly faster viral gene expression kinetics compared to epithelial-derived HeLa cells. In contrast, HeLa cells are human cervical cancer epithelial cells with relatively low efficiency in initial transcriptional activation of HSV-1 and higher baseline expression levels of endogenous antiviral factors (such as interferon-stimulated genes), which may lead to a marked delay in the expression of early immediate-early genes like ICP4.

(3) In responses 2-2, Fig 5K: An infection without transfection has not been included. This is important to understand kinetics of infection in transfected cells.

We sincerely appreciate the reviewer's insightful identification of this critical oversight. In all relevant experiments, we have strictly included empty vector transfection controls—serving as a baseline reference for each transfection group to eliminate potential influences from the transfection procedure itself and the vector background on viral replication, gene expression, and signaling pathways. The failure to clearly label this control in previous figure legends and main figures was indeed an omission in our presentation; we have now fully addressed this in the revised manuscript: all figures involving transfections (including Figures 3L, 3M, 5K, etc.) now explicitly indicate the "empty vector" control, and we have added detailed explanations in the figure legends and methods section regarding its role as an internal transfection control and procedural comparator. Once again, we thank the reviewer for their high level of professionalism in helping us enhance the completeness and scientific rigor of our data presentation.

(4) Why HDAC1 with deleted NES does not accumulate or looks like it is degraded? Why then ICP4 does not accumulate?

We sincerely apologize for the lack of clear labeling of the FLAG-HDAC1 ΔNES protein band in Author response image 2. This omission may have led reviewers to misinterpret its expression level as abnormal. After re-evaluation and improved annotation, Author response image 2 now clearly indicates the FLAG-HDAC1 ΔNES band its migration position corresponds to the expected molecular weight (slightly smaller than wild-type FLAG-HDAC1), and the band intensity is comparable to that of the empty vector and wild-type groups, indicating stable intracellular expression of this mutant protein without significant degradation. Therefore, its inhibitory effect on HSV-1 replication is not due to protein instability, but rather results from subcellular localization defects caused by the loss of nuclear export signal (NES): the ΔNES mutation causes HDAC1 to abnormally retain within the nucleus, ultimately leading to significant downregulation of ICP4 transcription and impaired protein accumulation.

Reviewer #2 (Public review):

Summary:

The authors discovered that HDAC1/2 are degraded in HSV-1 and PRV infections. They attempted to establish a new mechanism by which HDAC1/2 are translocated to the cytoplasm to be degraded in HSV-1 infection, and the degradation causes changes in histone acetylation to affect the DDR pathway.

Strengths:

(1) Interesting findings of HDAC1/2 degradation during HSV-1 and PRV infection, and it may impact more than the virology field.

(2) Significant work to identify the ubiquitin site in HDAC1/2 and K63 linkage.

Comments on revised version:

The authors added experiments to address the previous comments. The added knockdown and overexpression experiments provided sufficient support for the proposed mechanism. The conclusions are now strengthened. However, a few essential controls are still missing.

(1) Figure 3K: How does the expression level of Flag-HDAC1 variants compare to the endogenous HDAC1 level? The stripe probed by Flag antibody should be reprobed by HDAC1 antibody. Also, how does the K74R mutant affect histone acetylation? Moreover, the numbers between the panels are hard to read and have not been explained.

We sincerely thank the reviewers for their insightful and constructive feedback. In response to the comment on Figure 3K, we performed antibody re-probing of the Flag-immunoprecipitated or Flag-immunoblotted membranes with a validated HDAC1-specific antibody. Consistent with robust transfection and expression, both wild-type Flag-HDAC1 and its mutants including K74R exhibited markedly elevated total HDAC1 protein levels relative to vector control, confirming efficient exogenous expression and protein stability. To directly assess functional consequences, we evaluated global histone acetylation status in parallel samples and found that the K74R mutant induces significantly greater deacetylation than wild-type Flag-HDAC1, as demonstrated by pronounced reductions in H3K56ac and H4K8 acetylation levels. Finally, to improve clarity and readability, we have revised the lane annotations in Figure 3K—increasing font size, enhancing contrast, and ensuring consistent alignment—and fully documented these modifications in the updated figure legend.

(2) Figure 3M and 3L: DNA transfection per se frequently stimulates cell reactions that inhibit HSV-1 replication. Is the HSV-1 only sample transfected by empty vector or untransfected?

We sincerely appreciate the reviewer's insightful identification of this critical oversight. In all relevant experiments, we have strictly included empty vector transfection controls serving as a baseline reference for each transfection group to eliminate potential influences from the transfection procedure itself and the vector background on viral replication, gene expression, and signaling pathways. The failure to clearly label this control in previous figure legends and main figures was indeed an omission in our presentation; we have now fully addressed this in the revised manuscript: all figures involving transfections (including Figures 3L, 3M, 5K, etc.) now explicitly indicate the "empty vector" control, and we have added detailed explanations in the figure legends and methods section regarding its role as an internal transfection control and procedural comparator. Once again, we thank the reviewer for their high level of professionalism in helping us enhance the completeness and scientific rigor of our data presentation.

(3) Figure 4G-4J: What is the MDM2 knockdown efficiency?

During the construction of the MDM2 knockdown cell lines, we first systematically validated the knockdown efficiency by qRT-PCR. As shown in Figure 4A, compared to the control group (shCtrl), MDM2 mRNA levels were reduced by approximately 60% in shMDM2 cells, and protein expression also showed a corresponding significant decrease, confirming that the cell line had been successfully established and exhibited stable gene silencing effects.

(4) Figure 5F and line 400-401: "thereby preventing HDAC1 degradation-markedly impaired HSV-1 replication (Fig. 5F)." However, viral replication is not demonstrated in Figure 5F.

We sincerely appreciate the reviewer for pointing out the error in the figure legend numbering. Upon verification, the experimental data referred to in lines 400–401 of the original text and in Figure 5F actually correspond to the revised new Figure 5J. We apologize for failing to update the figure references in the main text during the revision process due to an oversight. We have now uniformly corrected all relevant descriptions in the text to "Figure 5J" and conducted a comprehensive review of all figure numbers, table numbers, and cross-references throughout the manuscript to confirm there are no other similar errors.

(5) Figure 5K: also need a control of empty vector. Furthermore, how does the HDAC1 ΔNES expression affect histone acetylation and DDR responses?

We sincerely thank the reviewers for their thoughtful and constructive feedback on Figure 5K. With regard to the empty vector control: all pertinent experiments in this study were performed with rigorous inclusion of an appropriate empty vector control (pCMV-Flag or its isogenic backbone), serving as the definitive negative control. The prior absence of this control in the figure representation was unintentional and reflects an oversight in data presentation—not in experimental design—and we offer our sincere apologies. We have now incorporated the empty vector control bands into Figure 5K and revised the figure legend to explicitly identify and describe this control. In addition, per the reviewers’ recommendation, we conducted a comprehensive assessment of HDAC1 ΔNES function, specifically examining its impact on global histone acetylation and canonical DNA damage response (DDR) activation. Quantitative immunoblotting and immunofluorescence analyses revealed that HDAC1 ΔNES expression leads to significantly greater reduction in H3K56ac and H4K8 acetylation compared with wild-type HDAC1. Moreover, upon induction of DNA damage, HDAC1 ΔNES-expressing cells exhibit attenuated DDR signaling, evidenced by diminished γH2AX focus formation, reduced CHK2 phosphorylation (p-CHK2), and blunted p53 stabilization and activation consistent with impaired DDR initiation or propagation (see Author response image 3). Collectively, these data indicate that nuclear retention of HDAC1 due to NES deletion not only potentiates its chromatin-targeted deacetylase activity but also contributes to suppression of DDR signaling, likely through epigenetic modulation of damage-sensing chromatin domains.

(6) Statements listed below are better moved to discussion after all data being presented. They are quite a stretch when looking at each figure by itself.

(i) Line 268-270: "Together, these findings indicate that HSV-1 selectively degrades class I HDACs, resulting in widespread histone hyperacetylation that fosters a chromatin state conducive to viral replication". ----may be okay for a statement.

(ii) Line 291-292: "providing initial evidence that HSV-1 infection promotes DDR activation through downregulation of HDAC1 expression"

(iii) Line 331-333: "Together, these results indicate that HSV-1 infection promotes K63-linked polyubiquitination of HDAC1/2 at conserved lysine residues, ultimately leading to their proteasomal degradation."

(iv) Line 334-336 is a repeated sentence.

We sincerely thank the reviewers for their thoughtful and constructive feedback. As noted, statements of mechanistic interpretation are not appropriate in the Results section; accordingly, we have relocated all such statements to the Discussion section. Furthermore, we have conducted a comprehensive line-by-line review of the manuscript to ensure that (i) every mechanistic inference is directly supported by experimental data presented in the Results, and (ii) integrative interpretations particularly those linking molecular observations to broader biological implications are confined exclusively to the Discussion.

Changes from the earlier version to the revised version

Abstract

Viral infection triggers a robust DNA damage response (DDR), reshaping the host chromatin landscape to facilitate viral replication. Here, we uncover a novel mechanism by which alphaherpesviruses exploit the DDR pathway. We demonstrated that herpes simplex virus 1 (HSV-1) and pseudorabies virus (PRV) induced selective degradation of class I histone deacetylases (HDAC1/2), leading to histone hyperacetylation and subsequent DDR activation. Strikingly, viral infection promoted nuclear export of HDAC1/2, followed by MDM2-mediated K63-linked polyubiquitination and proteasomal degradation in the cytoplasm. Pharmacological inhibition of either DDR signaling or HDAC1/2 nuclear export significantly affected viral replication in vitro and in vivo. Our findings reveal a unique viral strategy to hijack host epigenetic regulation for efficient replication and identify potential therapeutic targets for alphaherpesvirus infections.

Introduction

HSV-1 is a widespread neurotropic virus that infects more than two-thirds of the global population under the age of 50 [1]. After initial infection of mucosal epithelia and neuronal tissues, HSV-1 establishes lifelong latency within sensory ganglia [2]. Clinically, HSV-1 infection can present as oral or genital lesions and, in severe instances, may progress to herpes simplex encephalitis–a potentially fatal inflammation of the brain [3, 4]. Emerging evidence suggests a plausible association between HSV-1 and neurodegenerative conditions such as Alzheimer’s disease, linking chronic neuroinflammation and neuronal impairment to persistent viral infection [5, 6]. To sustain lifelong infection, HSV-1 employs sophisticated strategies to evade host immune defenses and maintain latent reservoirs.

Epigenetic regulation plays a critical role in governing HSV-1 latency and reactivation [7, 8]. During latency, the virus coopts host histone deacetylases (HDACs) to enforce viral genomic silencing through chromatin condensation [9]. Although HDAC1 and HDAC2 are central to chromatin stability and DNA damage response (DDR), their roles during alphaherpesvirus infection appear complex and context-dependent [10]. While these enzymes contribute to viral repression during latency [11], their potential involvement in promoting lytic replication–whether through degradation or functional inactivation–has remained inadequately explored.

Histone acetylation, a pivotal epigenetic modification that modulates chromatin accessibility and transcriptional activity, is dynamically regulated by histone acetyltransferases (HATs), which add acetyl groups to activate gene expression, and HDACs, which remove these groups to enforce repression [12, 13]. The HDAC family is categorized into four classes: Class I (HDAC1, 2, 3, 8), Class IIa/b, Class III (sirtuins), and Class IV (HDAC11) [14]. Among these, Class I HDACs–particularly HDAC1 and HDAC2–are indispensable for chromatin remodeling, genomic integrity, and immune modulation [15-18]. These enzymes contribute to host antiviral defense by fine-tuning the expression of immune-related genes. Nevertheless, the mechanisms through which HSV-1 overcomes this transcriptional repression remain incompletely understood [18].

Earlier version: HDAC1 serves as a key regulator of chromatin dynamics and innate immune signaling, interfacing with critical pathways such as NF-κB, JAK-STAT, and Toll-like receptor cascades to shape antiviral responses[19–22]. For example, in lung epithelial cells, HDAC1 facilitates STAT1 phosphorylation and enhances interferon-stimulated gene (ISG) activation, thereby restricting influenza A virus replication[23]. This underscores the dual function of HDAC1 in both epigenetic control and immune regulation. However, the specific strategies used by HSV-1 to manipulate HDAC1/2—particularly through ubiquitin-mediated degradation—to circumvent host immunity and enhance viral replication have not been clearly defined.

Revised version: HDAC1 serves as a key regulator of chromatin dynamics and innate immune signaling, interfacing with critical pathways such as NF-κB, JAK-STAT, and Toll-like receptor cascades to shape antiviral responses [19-22]. For example, in lung epithelial cells, HDAC1 facilitates STAT1 phosphorylation and enhances interferon-stimulated gene activation, thereby restricting influenza A virus replication [23]. This underscores the dual function of HDAC1 in both epigenetic control and immune regulation. However, the specific strategies used by HSV-1 to manipulate HDAC1/2–particularly through ubiquitin-mediated degradation–to circumvent host immunity and enhance viral replication have not been clearly defined.

Earlier version: This study reveals that HSV-1 induces the degradation of HDAC1/2 via an MDM2-dependent ubiquitination mechanism, resulting in elevated histone acetylation, chromatin relaxation, and activation of DNA damage response pathways that collectively enhance viral replication. These findings offer novel insights into the epigenetic subversion strategies employed by HSV-1 and underscore the therapeutic potential of targeting the HDAC–MDM2 axis in the treatment of herpesvirus infections.

Revised version: This study reveals that HSV-1 induces the degradation of HDAC1/2 via an MDM2-dependent ubiquitination mechanism, resulting in elevated histone acetylation, chromatin relaxation, and activation of DNA damage response pathways that collectively enhance viral replication. These findings offer novel insights into the epigenetic subversion strategies employed by HSV-1 and underscore the therapeutic potential of targeting the MDM2–HDAC axis in the treatment of herpesvirus infections.

Materials and Methods

Mice

Earlier version: Female C57BL/6J mice (6–8 weeks old) were purchased from the Experimental Animal Center of Zhengzhou University (Zhengzhou, China) and housed in specific-pathogen-free facilities under controlled environmental conditions, including a 12-hour light-dark cycle and a temperature of 22°C. For in vivo infection studies, mice were anesthetized with isoflurane prior to intranasal inoculation with HSV-1, at a dose of 1 × 10⁶ PFU per mouse in a total volume of 30 µL PBS. Clinical signs—including ruffled fur, hunched posture, reduced mobility, and body weight loss 15%—were monitored daily. On day 15 post-infection, mice were humanely euthanized by intravenous injection of sodium pentobarbital (90 mg/kg), following the “Guidelines for the Euthanasia of Laboratory Animals”. Immediately thereafter, tissues—including liver, spleen, lung, brain, and draining cervical lymph nodes—and serum were harvested under sterile conditions. All tissues were snap-frozen in liquid nitrogen within 60 seconds of excision and stored at −80°C until further analysis. Samples were processed for Western blotting. All procedures were in accordance.

Revised version: Female C57BL/6J mice (6–8 weeks old) were purchased from the Experimental Animal Center of Zhengzhou University (Zhengzhou, China) and housed in specific-pathogen-free facilities under controlled environmental conditions, including a 12-hour light-dark cycle and a temperature of 22°C. For in vivo infection studies, mice were anesthetized with isoflurane prior to intranasal inoculation with HSV-1, at a dose of 1 × 106 PFU per mouse in a total volume of 30 μL PBS. Clinical signs–including ruffled fur, hunched posture, reduced mobility, and body weight loss 15%–were monitored daily. On day 15 post-infection, mice were humanely euthanized by intravenous injection of sodium pentobarbital (90 mg/kg), following the “Guidelines for the Euthanasia of Laboratory Animals”. Immediately thereafter, tissues–including liver, spleen, lung, brain, and draining cervical lymph nodes–and serum were harvested under sterile conditions. All tissues were snap-frozen in liquid nitrogen within 60 seconds of excision and stored at −80°C until further analysis. Samples were processed for Western blotting. All procedures were in accordance with ethical guidelines.

Reagents and plasmids

Earlier version: Reagents were sourced as follows: Leptomycin B (HY-16909), MG-132 (HY-13259), 3-MA (HY-19312), cycloheximide (HY-12320), and chloroquine (HY-17589A) were purchased from MedChemExpress; berzosertib (S7102) was purchased from Selleck; DMSO (W387520) was purchased from Sigma-Aldrich. Unless otherwise specified, all inhibitors were added at 1 hour post-infection (hpi), after completion of viral adsorption and entry, to avoid interference with early viral processes. Antibodies against CHK1 (25887-1-AP), CHK2 (13954-1-AP), RAD51 (14961-1-AP), β-actin (66009-1-lg), P53 (10442-1-AP), HDAC1 (10197-1-AP), HDAC2 (12922-3-AP), HDAC4 (17449-1-AP), HDAC6 (12834-1-AP), HDAC11 (67949-1-Ig), and EGFP (50430-2-AP) were purchased from Proteintech; antibodies against p-P53 (9286), p-ATM (13050), ATM (2873), ATR (13934), p-ATR (2853), p-CHK1 (12302), p-CHK2 (2197), γ-H2AX (80312), H3 (4499), H4K8ac (2594), H4K12ac (13944), H4 (13919), H2AX (7631), H3K9ac (4658), H3K27ac (8173), UB (20326), and MDM2 (86934) were purchased from Cell Signaling Technology; H3K56ac antibody (07-677) was purchased from Millipore; ICP4 antibody (ab6514) was purchased from Abcam; ICP0 antibody (sc-53070) was purchased from Santa Cruz Biotechnology; FLAG antibody (F7425) was purchased from Sigma-Aldrich; HA antibody (A00169) was purchased from GenScript.

Revised version: Reagents were sourced as follows: Leptomycin B (HY-16909), MG-132 (HY-13259), 3-MA (HY-19312), cycloheximide (HY-12320), and chloroquine (HY-17589A) were purchased from MedChemExpress; berzosertib (S7102) was purchased from Selleck; DMSO (W387520) was purchased from Sigma-Aldrich. Unless otherwise specified, all inhibitors were added at 1-hour post-infection (hpi), after completion of viral adsorption and entry, to avoid interference with early viral processes. Antibodies against CHK1 (25887-1-AP), CHK2 (13954-1-AP), RAD51 (14961-1-AP), β-actin (66009-1-lg), P53 (10442-1-AP), HDAC1 (10197-1-AP), HDAC2 (12922-3-AP), HDAC4 (17449-1-AP), HDAC6 (12834-1-AP), HDAC11 (67949-1-Ig), and EGFP (50430-2-AP) were purchased from Proteintech; antibodies against p-P53 (9286), p-ATM (13050), ATM (2873), ATR (13934), p-ATR (2853), p-CHK1 (12302), p-CHK2 (2197), γ-H2AX (80312), H3 (4499), H4K8ac (2594), H4K12ac (13944), H4 (13919), H2AX (7631), H3K9ac (4658), H3K27ac (8173), UB (20326), and MDM2 (86934) were purchased from Cell Signaling Technology; H3K56ac antibody (07-677) was purchased from Millipore; ICP4 antibody (ab6514) was purchased from Abcam; ICP0 antibody (sc-53070) was purchased from Santa Cruz Biotechnology; FLAG antibody (F7425) was purchased from Sigma-Aldrich; HA antibody (A00169) was purchased from GenScript. HDAC1 and HDAC2 coding sequences were amplified from HEK293 cell cDNA and cloned into p3×FLAG-CMV-10 to generate FLAG-tagged expression constructs. MDM2 coding sequence was amplified and inserted into pEGFP-C1 to produce EGFP-MDM2. Site-directed mutagenesis was performed using the QuikChange Site-Directed Mutagenesis Kit (Agilent Technologies, 200523) per the manufacturer’s instructions to generate: FLAG-HDAC1 K8R, FLAG-HDAC1 K10R, FLAG-HDAC1 K31R, FLAG-HDAC1 K50R, FLAG-HDAC1 K58R, FLAG-HDAC1 K66R, FLAG-HDAC1 K74R, FLAG-HDAC1 K89R, FLAG-HDAC2 K75R, FLAG-HDAC1 (aa 1–329), FLAG-HDAC1 (aa 329–482), FLAG-HDAC1 (aa 1–100), FLAG-HDAC1 (aa 100–329), and EGFP-MDM2ΔRING (aa 1–435). FLAG-HDAC1 ΔNES (residues 207–216 deleted), in which the deletion encompasses the CRM1-dependent nuclear export signal as predicted by NESmapper. HA-tagged ubiquitin constructs [HA-UB (WT), HA-UB (K48), HA-UB (K63)] were provided by Dr. Bo Zhong (College of Life Sciences, Wuhan University, China).

Removed in the revised version: HDAC1 and HDAC2 coding sequences were amplified from HEK293 cell cDNA and cloned into p3×FLAG-CMV-10 to generate FLAG-tagged expression constructs. MDM2 coding sequence was amplified and inserted into pEGFP-C1 to produce EGFP-MDM2. Site-directed mutagenesis was performed using the QuikChange Site-Directed Mutagenesis Kit (Agilent Technologies, 200523) per the manufacturer’s instructions to generate: FLAG-HDAC1 K8R, FLAG-HDAC1 K10R, FLAG-HDAC1 K31R, FLAG-HDAC1 K50R, FLAG-HDAC1 K58R, FLAG-HDAC1 K66R, FLAG-HDAC1 K74R, FLAG-HDAC1 K89R, FLAG-HDAC2 K75R, FLAG-HDAC1 (aa 1–329), FLAG-HDAC1 (aa 329–482), FLAG-HDAC1 (aa 1–100), FLAG-HDAC1 (aa 100–329),, and EGFP-MDM2ΔRING (aa 1–435). FLAG-HDAC1 ΔNES (residues 207–216 deleted), in which the deletion encompasses the CRM1-dependent nuclear export signal as predicted by NESmapper. HA-tagged ubiquitin constructs [HA-UB (WT), HA-UB (K48), HA-UB (K63)] were provided by Dr. Bo Zhong (College of Life Sciences, Wuhan University, China).

Cells

Earlier version: Cell lines (HeLa/ATCC CL-82, Vero/ATCC CL-81, 3D4/21/ATCC CRL-2843, HEK293T/ATCC CRL-11268) were cultured in DMEM or RPMI 1640 supplemented with 10% fetal bovine serum (FBS), 1% penicillin/streptomycin, and 1% L-glutamine at 37°C under 5% CO₂ in a humidified incubator. For transient transfections, DNA constructs were delivered to HeLa and HEK293T cells using Lipofectamine 3000 (Invitrogen, L3000001) per the manufacturer’s protocol.

Revised version: Cell lines (HeLa/ATCC CL-82, Vero/ATCC CL-81, 3D4/21/ATCC CRL-2843, HEK293T/ATCC CRL-11268) were cultured in DMEM 10566–016 or RPMI 1640 (Gibco, 11965092) supplemented with 10% fetal bovine serum (FBS; Gibco, 10099141C), 1% penicillin/streptomycin and 1% L-glutamine (B540732, Sangon) at 37°C under 5% CO2 in a humidified incubator. For transient transfections, DNA constructs were delivered to HeLa and HEK293T cells using Lipofectamine 3000 (Invitrogen, L3000001) per the manufacturer’s protocol.

To generate shRNA-mediated knockdown cell lines, HEK293T cells were co-transfected with control or gene-specific shRNA (Table S1), packaging plasmid psPAX2, and envelope plasmid pMD2.G. Post-transfection (6 h), medium was replaced, and lentiviral particles were harvested at 48 hours. Parental cells were infected with viral supernatant and selected with puromycin to establish stable knockdown lines.

To generate siRNA-mediated knockdown cell lines, HeLa cells were transfected with either a non-targeting control siRNA or gene-specific siRNAs (Table S1) using Lipofectamine™ RNAiMAX (Invitrogen, 13778030). Transfection was performed at a confluency of approximately 30%, and target gene and protein expression levels were assessed 48 hours post-transfection.

Viruses

Earlier version: HSV-1 strain F (provided by Dr. Chun-Fu Zheng, University of Calgary, Canada) and PRV-QXX were propagated per established protocols [24]. Viral titers were determined by plaque assays in Vero cells. For in vitro infections, cells were infected with HSV-1 or PRV-QXX at an MOI of 1. For in vivo studies, mice were intranasally inoculated with HSV-1 (1 × 10⁶ pfu/mouse).

Revised version: HSV-1 strain F (provided by Dr. Chun-Fu Zheng, University of Calgary, Canada) and PRV-QXX were propagated per established protocols [24]. Viral titers were determined by plaque assays in Vero cells. For in vitro infections, cells were infected with HSV-1 or PRV-QXX at an MOI of 1. For in vivo studies, mice were intranasally inoculated with HSV-1 (1 × 106 pfu/mouse).

Immunoblotting, ubiquitination and Co-immunoprecipitation (Co-IP)

Earlier version: For immunoblotting, cells were lysed in RIPA buffer [50 mM Tris-HCl (pH 8.0), 150 mM NaCl, 1% Triton X-100, 1% sodium deoxycholate, 0.1% SDS, 2 mM MgCl₂] supplemented with protease/phosphatase inhibitors. Proteins were resolved by SDS-PAGE, transferred to PVDF membranes, and blocked with 5% nonfat milk in TBST (1 hour, RT). Membranes were incubated with primary antibodies (overnight, 4°C), then HRP-conjugated secondary antibodies (1 hour, RT). Signals were detected using Laminate Crescendo Western HRP Substrate (Millipore, Cat. No. WBLUR0500) on a GE AI600 imager.

Revised version: For immunoblotting, cells were lysed in RIPA buffer [50 mM Tris-HCl (pH 8.0), 150 mM NaCl, 1% Triton X-100, 1% sodium deoxycholate, 0.1% SDS, 2 mM MgCl2] supplemented with protease/phosphatase inhibitors. Proteins were resolved by SDS-PAGE, transferred to PVDF membranes, and blocked with 5% nonfat milk in TBST (1 hour, RT). Membranes were incubated with primary antibodies (overnight, 4°C), then HRP-conjugated secondary antibodies (1 hour, RT). Signals were detected using Laminate Crescendo Western HRP Substrate (Millipore, WBLUR0500) on a GE AI600 imager.

For ubiquitination assay, cells were lysed in IP buffer [50 mM Tris-HCl (pH 7.4), 150 mM NaCl, 1% NP-40, 1% sodium deoxycholate, 5 mM EDTA, 5 mM EGTA] and clarified (16,000 × g, 10 minutes, 4°C). Subsequently, 900 μL aliquots were incubated with 40 μL of a 1:1 slurry of Sepharose beads conjugated to anti-HDAC1 or anti-FLAG mouse monoclonal antibodies (4°C, 4 hours). Beads were washed four times with IP buffer, eluted in SDS sample buffer (10 minutes, boiling), and analyzed by immunoblotting.

Earlier version: For Co-IP assay, cells were lysed in Co-IP buffer [50 mM Tris-HCl (pH 7.4), 150 mM NaCl, 1% NP-40, 5 mM EDTA, 5 mM EGTA] and clarified (16,000 × g, 10 minutes, 4°C). Aliquots (900 μl) were incubated with 40 μl of a 1:1 slurry of Sepharose beads conjugated to IgG (GE Healthcare) or anti-FLAG mouse monoclonal antibody (4°C, 4 hours). Beads were washed three times with Co-IP buffer, eluted in SDS sample buffer (10 minutes, boiling), and analyzed by immunoblotting.

Revised version: For Co-IP assay, cells were lysed in Co-IP buffer [50 mM Tris-HCl (pH 7.4), 150 mM NaCl, 1% NP-40, 5 mM EDTA, 5 mM EGTA] and clarified (16,000 × g, 10 minutes, 4°C). Aliquots (900 μL) were incubated with 40 μL of a 1:1 slurry of Sepharose beads conjugated to IgG (GE Healthcare, AI600) or anti-FLAG mouse monoclonal antibody (4°C, 4 hours). Beads were washed three times with Co-IP buffer, eluted in SDS sample buffer (10 minutes, boiling), and analyzed by immunoblotting.

Immunofluorescence assay

Earlier version: Cells were fixed in 4% paraformaldehyde at room temperature for 20 min on coverslips in 12-well plates. Following three washes with PBS, cells were permeabilized with 0.2% Triton X-100 for 20 min and subsequently blocked with 10% FBS. Specific primary antibodies, diluted in 10% FBS, were then applied to the cells and incubated for 1 h at room temperature. After three PBS washes, cells were exposed to the appropriate secondary antibodies, also diluted in 10% FBS, for 1 h at room temperature. Nuclei were stained with DAPI for 5 min at room temperature, mounted using Prolong Diamond (Invitrogen), and visualized using a Zeiss LSM 800 confocal microscope.

Revised version: Cells were fixed in 4% paraformaldehyde at room temperature for 20 min on coverslips in 12-well plates. Following three washes with PBS, cells were permeabilized with 0.2% Triton X-100 for 20 min and subsequently blocked with 10% FBS. Specific primary antibodies, diluted in 10% FBS, were then applied to the cells and incubated for 1 h at room temperature. After three PBS washes, cells were exposed to the appropriate secondary antibodies, also diluted in 10% FBS, for 1 h at room temperature. Nuclei were stained with DAPI for 5 min at room temperature, mounted using Prolong Diamond (Invitrogen, P36970), and visualized using a Zeiss LSM 800 confocal microscope.

Comet assay

Earlier version: Cells were seeded in 6-well plates and treated accordingly. Frosted microscope slides were coated with 0.5% normal melting point agarose. A mixture of 10 μL DMEM containing around 10,000 cells and 75 μL of 0.7% low melting point agarose was layered onto the pre-coated agarose. An additional 75 μL of 0.7% low melting point agarose formed a third layer. The slides were lysed in a buffer containing 10 mM Tris-HCl, pH 10.0, 2.5 M NaCl, 100 mM Na2EDTA, 1% Triton X-100, and 10% DMSO for 2 hours at 4°C. Post-lysis, the slides were treated with an electrophoresis solution (300 mM NaOH, 1 mM Na2EDTA, pH >13) for 40 minutes, electrophoresed at 20 V (∼300 mA) for 25 minutes, and neutralized with 0.4 mM Tris-HCl (pH 7.5). Subsequently, cells were stained with propidium iodide (PI, 5 μg/mL) and examined using a Zeiss LSM 800 confocal microscope. DNA damage was assessed based on tail moment using CometScore software.

Revised version: Cells were seeded in 6-well plates and treated accordingly. Frosted microscope slides were coated with 0.5% normal melting point agarose. A mixture of 10 μL DMEM containing around 10,000 cells and 75 μL of 0.7% low melting point agarose was layered onto the pre-coated agarose. An additional 75 μL of 0.7% low melting point agarose formed a third layer. The slides were lysed in a buffer containing 10 mM Tris-HCl, pH 10.0, 2.5 M NaCl, 100 mM Na2EDTA, 1% Triton X-100, and 10% DMSO for 2 hours at 4°C. Post-lysis, the slides were treated with an electrophoresis solution (300 mM NaOH, 1 mM Na2EDTA, pH >13) for 40 minutes, electrophoresed at 20 V (~300 mA) for 25 minutes, and neutralized with 0.4 mM Tris-HCl (pH 7.5). Subsequently, cells were stained with propidium iodide (5 μg/mL) and examined using a Zeiss LSM 800 confocal microscope. DNA damage was assessed based on tail moment using CometScore software.

qRT-PCR analysis

Earlier version: Total RNA was isolated using TRIzol reagent (TaKaRa) and reverse-transcribed into cDNA with the PrimeScript RT reagent kit (TaKaRa). qRT-PCR was performed in triplicate using SYBR Premix Ex Taq (TaKaRa) according to the manufacturer’s protocol. Expression levels were normalized to β-actin as an internal reference. Melting curve analysis confirmed amplification specificity by verifying single-product formation in all reactions. Relative gene expression was calculated using the 2−ΔΔCt method. Primer sequences are provided in Table S1

Revised version: Total RNA was isolated using TRIzol reagent (TaKaRa, 9108) and reverse-transcribed into cDNA with the PrimeScript RT reagent kit (TaKaRa, RR047A). qRT-PCR was performed in triplicate using SYBR Premix Ex Taq (TaKaRa, RR820A) according to the manufacturer’s protocol. Expression levels were normalized to β-actin as an internal reference. Melting curve analysis confirmed amplification specificity by verifying single-product formation in all reactions. Relative gene expression was calculated using the 2−ΔΔCt method. Primer sequences are provided in Table S1.

Plaque assay

Earlier version: Vero cells (seeded in six-well plates) were grown to confluence, infected with serially diluted viruses (10⁻¹–10⁻⁷; 1 hour, 37°C), and washed with PBS to remove residual inoculum. DMEM containing 1% methylcellulose (4 mL/well) was added, and cells were incubated for 4–5 days. Post-incubation, cells were fixed with 4% paraformaldehyde (15 minutes), stained with 1% crystal violet (30 minutes), and plaques were quantified.

Revised version: Vero cells (seeded in six-well plates) were grown to confluence, infected with serially diluted viruses (10−1–10−7; 1 hour, 37°C), and washed with PBS to remove residual inoculum. DMEM containing 1% methylcellulose (4 mL/well) was added, and cells were incubated for 4–5 days. Post-incubation, cells were fixed with 4% paraformaldehyde (15 minutes), stained with 1% crystal violet (30 minutes), and plaques were quantified.

Statistical analysis

Statistical analyses were performed using GraphPad Prism 8 software. Comparisons between two groups were evaluated using a two-tailed Student’s t-test. Significance was defined as P < 0.05. Data are presented as mean ± SD of three independent experiments. Kaplan-Meier survival curves were generated and analyzed for mouse survival assessment.

Results

HSV-1 infection promotes HDAC1/2 degradation and histone hyperacetylation

Post-translational modifications of histones play a critical role in regulating chromatin remodeling and gene expression. HDACs, a highly conserved enzyme family [25], catalyze the removal of acetyl groups from histones, thereby promoting chromatin condensation and transcriptional repression. To investigate the effect of HSV-1 infection on HDAC expression and histone acetylation, we examined the protein levels of key HDAC isoforms following viral infection. Immunoblotting analysis revealed a marked decrease in HDAC1 and HDAC2 protein levels in cells infected with either HSV-1 or PRV, whereas the expression of HDAC4, HDAC6, and HDAC11 remained unchanged (Fig. 1A, B). In contrast, qRT-PCR results showed no alterations in HDAC1 and HDAC2 mRNA levels (Fig. 1C, D), suggesting that their depletion is regulated at the post-translational level.

Viral infection induces degradation of class I HDACs and hyperacetylation of histones H3 and H4.

Earlier version: (A) Immunoblotting analysis of the indicated HDACs in HeLa cells infected with HSV-1 (MOI = 1) for the indicated time points. (B) Immunoblotting analysis of the indicated HDACs in 3D4/21 cells infected with PRV-QXX (MOI = 1) for the indicated time points. (C and D) qRT-PCR analysis of HDAC1 and HDAC2 mRNA levels, normalized to β-actin expression, in HeLa cells infected with HSV-1 (C) or 3D4/21 cells infected with PRV-QXX (D) (MOI = 1). Data are mean ± SEM; n = 3. **P < 0.01, ***P < 0.001. (E) Immunoblotting of the indicated proteins in HeLa cells from infected with HSV-1 (MOI = 1) for the indicated time points. (F) Immunoblotting of the indicated proteins in 3D4/21 cells from infected with PRV-QXX (MOI = 1) for the indicated time points. (G) Immunoblotting analysis of the indicated proteins in liver tissues from mice mock-infected or infected with HSV-1 (1 × 10⁶ pfu per mouse) at 5 days post-infection (n = 3). (H) Immunoblotting of the indicated proteins in HeLa cells transfected with siControl or siHDAC1/2.

Revised version: (A) Immunoblotting analysis of the indicated HDACs in HeLa cells infected with HSV-1 (MOI = 1) for the indicated time points. (B) Immunoblotting analysis of the indicated HDACs in 3D4/21 cells infected with PRV-QXX (MOI = 1) for the indicated time points. (C and D) qRT-PCR analysis of HDAC1 and HDAC2 mRNA levels, normalized to β-actin expression, in HeLa cells infected with HSV-1 (C) or 3D4/21 cells infected with PRV-QXX (D) (MOI = 1). Data are mean ± SEM; n = 3. **P < 0.01, ***P < 0.001. (E) Immunoblotting of the indicated proteins in HeLa cells infected with HSV-1 (MOI = 1) for the indicated time points. (F) Immunoblotting of the indicated proteins in 3D4/21 cells infected with PRV-QXX (MOI = 1) for the indicated time points. (G) Immunoblotting analysis of the indicated proteins in liver tissues from mice mock-infected or infected with HSV-1 (1 × 106 pfu per mouse) at 5 days post-infection (n = 3). (H) Immunoblotting of the indicated proteins in HeLa cells transfected with siControl or siHDAC1/2. Note: The numerical values beneath each lane in the Western blot images represent the relative abundance of the corresponding protein bands, as quantified by densitometric analysis and normalized to the signal intensity of the first lane, which serves as the control group.

Earlier version: To assess the functional impact of HDAC1/2 reduction, we examined global histone acetylation patterns. Site-specific immunoblotting demonstrated elevated acetylation at multiple lysine residues, including H3K9, H3K27, H3K56, H4K8, and H4K12, in HSV-1-infected HeLa cells (Fig. 1E). A similar hyperacetylation profile was observed in PRV-infected 3D4/21 cells and in murine liver tissues infected with HSV-1 (Fig. 1F, G). Next, we assessed the effect of HDAC1/2 knockdown on global histone acetylation and observed significantly elevated acetylation at multiple evolutionarily conserved lysine residues—including H3K9, H3K27, H3K56, H4K8, and H4K12—consistent with genome-wide histone hyperacetylation (Fig. H). Together, these findings indicate that HSV-1 selectively degrades class I HDACs, resulting in widespread histone hyperacetylation that fosters a chromatin state conducive to viral replication.

Revised version: To assess the functional impact of HDAC1/2 reduction, we examined global histone acetylation patterns. Site-specific immunoblotting demonstrated elevated acetylation at multiple lysine residues, including H3K9, H3K27, H3K56, H4K8, and H4K12, in HSV-1-infected HeLa cells (Fig. 1E). A similar hyperacetylation profile was observed in PRV-infected 3D4/21 cells and in murine liver tissues infected with HSV-1 (Fig. 1F, G). Next, we assessed the effect of HDAC1/2 knockdown on global histone acetylation and observed significantly elevated acetylation at multiple evolutionarily conserved lysine residues–including H3K9, H3K27, H3K56, H4K8, and H4K12–consistent with genome-wide histone hyperacetylation (Fig. 1H). Collectively, these orthogonal lines of evidence—viral infection–induced degradation, cross-species conservation, and knockdown-mediated phenocopy—support the conclusion that HSV-1–mediated depletion of class I HDACs drives genome-wide histone hyperacetylation, establishing a chromatin environment permissive for viral transcription and replication.

HSV-1 infection induces DDR through HDAC1/2 depletion and histone hyperacetylation

Earlier version: Given that excessive histone acetylation can destabilize chromatin structure, we sought to elucidate the mechanism by which this epigenetic modification activates the DDR pathway. Accumulating evidence underscores the pivotal role of histone modifications in maintaining genomic stability, with dynamic changes in chromatin architecture being intimately associated with DDR initiation. To determine whether HSV-1 infection induces DDR activation, we evaluated the phosphorylation of γ-H2AX, a well-established marker of DNA double-strand breaks. Immunofluorescence microscopy revealed a significant increase in γ-H2AX foci in cells infected with HSV-1 or PRV (Fig. 2A-C). Comet assays (single-cell gel electrophoresis) quantify DNA fragmentation, with an increased tail moment reflecting elevated levels of DNA damage—findings that further corroborate enhanced DNA fragmentation under these experimental conditions (Fig. 2D–G).Given the central role of ATM/ATR signaling in the DDR, we proceeded to examine the activation of key checkpoint kinases. Immunoblotting analysis confirmed phosphorylation of ATM, ATR, Chk1, Chk2, H2AX, and RAD51 in HSV-1-infected HeLa cells and 3D4/21 cells (Fig. 2H, I). Subsequently, we assessed the effect of HDAC1 knockdown on the DDR pathway and found that HDAC1-depleted HeLa cells exhibited reduced phosphorylation of ATM, ATR, and H2AX(Fig. 2J), providing initial evidence that HSV-1 infection promotes DDR activation through downregulation of HDAC1 expression.

Revised version: Given that excessive histone acetylation can destabilize chromatin structure, we sought to elucidate the mechanism by which this epigenetic modification activates the DDR pathway. Accumulating evidence underscores the pivotal role of histone modifications in maintaining genomic stability, with dynamic changes in chromatin architecture being intimately associated with DDR initiation. To determine whether HSV-1 infection induces DDR activation, we evaluated the phosphorylation of γ-H2AX, a well-established marker of DNA double-strand breaks. Immunofluorescence microscopy revealed a significant increase in γ-H2AX foci in cells infected with HSV-1 or PRV (Fig. 2A–C). Comet assays (single-cell gel electrophoresis) quantify DNA fragmentation, with an increased tail moment reflecting elevated levels of DNA damage–findings that further corroborate enhanced DNA fragmentation under these experimental conditions (Fig. 2D–G). Given the central role of ATM/ATR signaling in the DDR, we proceeded to examine the activation of key checkpoint kinases. Immunoblotting analysis confirmed phosphorylation of ATM, ATR, Chk1, Chk2, H2AX, and RAD51 in HSV-1-infected HeLa cells and 3D4/21 cells (Fig. 2H, I). Subsequently, we assessed the effect of HDAC1 knockdown on the DDR pathway and found that HDAC1-depleted HeLa cells exhibited reduced phosphorylation of ATM,ATR, and H2AX (Fig. 2J), indicating that HDAC1 depletion is sufficient to potentiate DDR signaling. Together, these data demonstrate that HSV-1 infection induces a coordinated, ATM/ATR-dependent DDR, and that HDAC1 loss represents a functionally relevant upstream event contributing to this response.

Viral infection triggers chromatin dysfunction and DDR activation via HDAC degradation.

Earlier version: (A, B and C) Immunofluorescence staining of γ-H2AX (red) and H2AX (green) in HSV-1-infected HeLa cells (A) or PRV-infected 3D4/21 cells (B) (MOI = 1). Panel C presents the quantitative analysis of γ-H2AX foci intensity or nuclear fluorescence signal derived from panels A and B. Scale bars: 10 μm. Data are mean ± SEM; n = 100 cells/group. ***P < 0.001. (D and E) Comet assay assessing DNA damage in HeLa cells infected with HSV-1 (MOI = 1). Data are mean ± SEM; n = 40 cells/group. **P < 0.01, ***P < 0.001. (F and G) Comet assay assessing DNA damage in 3D4/21 cells infected with PRV (MOI = 1) at the indicated time points. Data are mean ± SEM; n = 40 cells/group. **P < 0.01, ***P < 0.001. (H and I) Immunoblotting analysis of DDR markers (p-ATM, p-ATR, p-Chk1, p-Chk2, γ-H2AX) and viral ICP4/EP0 in HSV-1-infected HeLa cells (G) or PRV-QXX-infected 3D4/21 cells (H) (MOI = 1). (J) Immunoblotting analysis of DDR markers (p-ATM, p-ATR, γ-H2AX) and HDAC1/HDAC2 in HeLa cells transfected with siHDAC1/2. (K) HeLa cells were treated with vehicle and Berzosertib (0-60 μM) for 0–48 h. Cell proliferation was analyzed by CCK-8 assay. ns, no significance. (L) Viral titers in HSV-1-infected HeLa cells (MOI = 2) treated with berzosertib (50 nM) or vehicle. Data are mean ± SEM; n = 3. *P < 0.05, **P < 0.01, ***P < 0.001 (M) Survival curves of HSV-1-infected mice (1 × 10⁶ PFU/mouse) treated with berzosertib (20 mg/kg) or vehicle (n = 12/group), ***P < 0.001.

Revised version: (A–C) Immunofluorescence staining of γ-H2AX (red) and H2AX (green) in HSV-1-infected HeLa cells (A) or PRV-infected 3D4/21 cells (B) (MOI = 1). Panel C presents the quantitative analysis of γ-H2AX foci intensity or nuclear fluorescence signal derived from panels A and B. Scale bars: 10 μm. Data are mean ± SEM; n = 100 cells/group. ***P < 0.001. (D and E) Comet assay assessing DNA damage in HeLa cells infected with HSV-1 (MOI = 1). Data are mean ± SEM; n = 40 cells/group. **P < 0.01, ***P < 0.001. (F and G) Comet assay assessing DNA damage in 3D4/21 cells infected with PRV (MOI = 1) at the indicated time points. Data are mean ± SEM; n = 40 cells/group. **P < 0.01, ***P < 0.001. (H and I) Immunoblotting analysis of DDR markers (p-ATM, p-ATR, p-Chk1, p-Chk2, γ-H2AX) and viral ICP4/EP0 in HSV-1-infected HeLa cells (G) or PRV-QXX-infected 3D4/21 cells (H) (MOI = 1). (J) Immunoblotting analysis of DDR markers (p-ATM, p-ATR, γ-H2AX) and HDAC1/HDAC2 in HeLa cells transfected with siHDAC1/2. (K) HeLa cells were treated with vehicle and Berzosertib (0–60 μM) for 0–48 h. Cell proliferation was analyzed by CCK-8 assay. ns, no significance. (L) Viral titers in HSV-1-infected HeLa cells (MOI = 2) treated with berzosertib (50 nM) or vehicle. t = 0 h defined as time of medium replacement post-adsorption.Data are mean ± SEM; n = 3. *P < 0.05, **P < 0.01, ***P < 0.001 (M) Survival curves of HSV-1-infected mice (1 × 106 PFU/mouse) treated with berzosertib (20 mg/kg) or vehicle (n = 12/group), ***P < 0.001. Note: The numerical values beneath each lane in the Western blot images represent the relative abundance of the corresponding protein bands, as quantified by densitometric analysis and normalized to the signal intensity of the first lane, which serves as the control group.

Next, we pharmacologically validated the functional relevance of ATR in HSV-1 replication using Berzosertib, a selective and clinically advanced ATR inhibitor. CCK-8 assays confirmed that the concentration of Berzosertib used (100 nM) did not significantly affect HeLa cell viability or proliferation over 48 hours (Fig. 2K), yet it markedly reduced viral titers in plaque assays (Fig. 2L). Collectively, these findings demonstrate that HSV-1 infection activates the DDR, likely through degradation of HDAC1 and HDAC2 and the resulting histone hyperacetylation. This virus-triggered DDR activation highlights the profound disruption of host chromatin homeostasis and genomic integrity by HSV-1.

HSV-1 promotes the ubiquitin-proteasome degradation of HDAC1/2

To elucidate the mechanism responsible for HDAC1/2 degradation, we focused on their post-translational regulation. Since HDAC1/2 mRNA levels were unchanged after infection, we hypothesized that the reduction in HDAC1/2 protein is mediated through enhanced proteolytic degradation. To identify the relevant degradation pathway, HSV-1-infected HeLa cells were treated with either the proteasome inhibitor MG-132 or autophagy inhibitors (3-MA and chloroquine). Immunoblotting analysis indicated that only MG-132 rescued HDAC1/2 protein levels (Fig. 3A), confirming that degradation occurs primarily via the proteasomal pathway.

Viral infection induces K63-linked ubiquitination of HDAC1/2.

Earlier version: (A) Immunoblotting analysis of the indicated proteins in HeLa cells either mock-infected or infected with HSV-1 (MOI = 1), treated with vehicle, 3-MA (10 μM), MG-132 (10 μM), or chloroquine (10 μM) for the indicated times. (B) Ubiquitination assay of HDAC1 in HeLa cells either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. (C) Ubiquitination assay of FLAG-HDAC1 in HeLa cells transfected with HA-ubiquitin variants (WT, K48, K63), and either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. (D) Schematic representation of HDAC1 deletion mutants. (E-I) Ubiquitination assays of FLAG-HDAC1 mutant variants in HeLa cells either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. (J) Ubiquitination assays of FLAG-HDAC2 variants (WT and K75R) in HeLa cells either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. (K) Immunoblotting analysis of the indicated proteins in HeLa cells transfected with FLAG-HDAC1 and FLAG-HDAC1 (K74R), and either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. (L) qRT-PCR analysis of HSV-1 ICP0,ICP8 and gB mRNA levels, normalized to β-actin expression, in HeLa cells transfected with FLAG-HDAC1 and FLAG-HDAC1 (K74R), and either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. Data are mean ± SEM; n = 3. ns > 0.05,*P < 0.05,**P < 0.01, ***P < 0.001. (M) Viral titer analysis in HeLa cells transfected with FLAG-HDAC1 or FLAG-HDAC1 (K74R), and either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. Data are mean ± SEM; n = 3. **P < 0.01, ***P < 0.001.

Revised version: (A) Immunoblotting analysis of the indicated proteins in HeLa cells either mock-infected or infected with HSV-1 (MOI = 1), treated with vehicle, 3-MA (10 μM), MG-132 (10 μM), or chloroquine (10 μM) for the indicated times. (B) Ubiquitination assay of HDAC1 in HeLa cells either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. (C) Ubiquitination assay of FLAG-HDAC1 in HeLa cells transfected with HA-ubiquitin variants (WT, K48, K63), and either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. (D) Schematic representation of HDAC1 deletion mutants. (E–I) Ubiquitination assays of FLAG-HDAC1 mutant variants in HeLa cells either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. (J) Ubiquitination assays of FLAG-HDAC2 variants (WT and K75R) in HeLa cells either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. (K) Immunoblotting analysis of the indicated proteins in HeLa cells transfected with empty vector, FLAG-HDAC1 (wild-type), or FLAG-HDAC1 (K74R), followed by mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. (L) qRT-PCR analysis of HSV-1 ICP0, ICP8 and gB mRNA levels, normalized to β-actin expression, in HeLa cells transfected with empty vector, FLAG-HDAC1 (wild-type), or FLAG-HDAC1 (K74R), followed by mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. Data are mean ± SEM; n = 3. ns > 0.05, *P < 0.05, **P < 0.01,***P < 0.001. (M) Viral titer analysis in HeLa cells transfected with FLAG-HDAC1 or FLAG-HDAC1 (K74R), and either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. Data are mean ± SEM; n = 3. **P < 0.01, ***P < 0.001. Note: The numerical values beneath each lane in the Western blot images represent the relative abundance of the corresponding protein bands, as quantified by densitometric analysis and normalized to the signal intensity of the first lane, which serves as the control group.

To examine whether HSV-1 infection induced ubiquitination of HDAC1, endogenous immunoprecipitation assays were performed, which revealed increased ubiquitination of HDAC1 following infection (Fig. 3B). Subsequent ubiquitination assays using wild-type (WT), K48-only, and K63-only ubiquitin mutants demonstrated that HDAC1 undergoes K63-linked–but not K48-linked–polyubiquitination (Fig. 3C). Using a series of HDAC1 truncation mutants, we mapped the ubiquitination site to the N-terminal region (amino acids 1–100) (Fig. 3D–F). Site-directed mutagenesis further identified lysine 74 (K74) as the critical residue required for ubiquitination (Fig. 3G–I). Sequence alignment showed that HDAC2 K75 corresponds to HDAC1 K74 (Fig. 3J), and a K75R mutation in HDAC2 similarly abolished ubiquitination (Fig. 3J).

Earlier version: To further validate the role of HDAC1 K74, we first overexpressed HDAC1 WT and HDAC1 K74R and assessed phosphorylation of ATM, ATR, and H2AX; results showed that HDAC1 K74R overexpression suppressed HSV-1–induced DDR signaling(Fig. 3K). Furthermore, RT–PCR analysis revealed that overexpression of both HDAC1 WT and HDAC1 K74R inhibited HSV-1 ICP0,ICP8 and gB expression, but HDAC1 WT exerted a significantly stronger inhibitory effect (Fig. 3L). Consistently, HDAC1 K74R overexpression reduced viral titers in plaque assays (Fig. 3M). Together, these results indicate that HSV-1 infection promotes K63-linked polyubiquitination of HDAC1/2 at conserved lysine residues, ultimately leading to their proteasomal degradation. Together, these results indicate that HSV-1 infection promotes K63-linked polyubiquitination of HDAC1/2 at conserved lysine residues, ultimately leading to their proteasomal degradation.

Revised version: To further validate the role of HDAC1 K74, we first overexpressed HDAC1 WT and HDAC1 K74R and assessed phosphorylation of ATM, ATR, and H2AX; results showed that HDAC1 K74R overexpression suppressed HSV-1–induced DDR signaling (Fig. 3K). Furthermore, qRT-PCR analysis revealed that overexpression of both HDAC1 WT and HDAC1 K74R inhibited HSV-1 ICP0, ICP8 and gB expression, but HDAC1 WT exerted a significantly stronger inhibitory effect (Fig. 3L). Consistently, HDAC1 K74R overexpression reduced viral titers in plaque assays (Fig. 3M). Collectively, these data demonstrate that HDAC1 K74 acetylation status modulates its capacity to restrain HSV-1 replication, likely by fine-tuning its regulatory influence on DDR signaling and viral gene expression

MDM2 functions as the E3 ligase mediating HDAC1/2 ubiquitination and degradation

Having established the occurrence of K63-linked ubiquitination (Fig. 3), we next sought to identify the E3 ubiquitin ligase responsible. An RNAi screen targeting candidate E3 ligases revealed MDM2 as the principal mediator of HDAC1/2 degradation. Knockdown of MDM2–but not of PIRH2, KCTD11, CHFR, UHRF1, or TRIM46–effectively prevented HDAC1/2 depletion (Fig. 4A, B). Co-IP assay further confirmed an enhanced interaction between HDAC1 and MDM2 following HSV-1 infection (Fig. 4C).

MDM2 mediates viral-induced HDAC1/2 degradation via K74/K75 ubiquitination.

Earlier version: (A) qRT-PCR analysis validation of E3 ligase knockdown efficiency (shPIRH2, shKCTD11, shMDM2, shCHFR, shUHRF1, shTRIM46) in HeLa cells. Data are mean ± SEM; n = 3. ***P < 0.001. (B) Immunoblotting analysis of HDAC1/2 stability in E3 ligase-knockdown HeLa cells infected with HSV-1 (MOI = 1). (C) Co-IP assay of endogenous HDAC1-MDM2 interaction in HSV-1-infected HeLa cells (MOI = 1; 24 hpi). (D) Co-IP assay of FLAG-HDAC1 with EGFP-MDM2 (WT or ΔRING) in HSV-1-infected HeLa cells (MOI = 1; 24 hpi). (E) Co-IP assay of FLAG-HDAC1 (WT or K74R) with EGFP-MDM2 in HSV-1-infected HeLa cells (MOI = 1; 24 hpi). (F) Ubiquitination assays of FLAG-HDAC1 (WT or K74R) co-expressed with EGFP-MDM2 (WT or ΔRING) in HSV-1-infected HeLa cells (MOI = 1; 24 hpi). (G-I) qRT-PCR analysis of HSV-1 ICP0 (H),ICP8 (I)and gB (J) mRNA levels, normalized to β-actin expression, in Scramble and shMDM2 HeLa cells, and either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. Data are mean ± SEM; n = 3. ns > 0.05,*P < 0.05,**P < 0.01, ***P < 0.001. (J) Viral titer analysis in Scramble and shMDM2 HeLa cells, either mock-infected or infected with HSV-1 (MOI = 2) at 24 hpi. Data are mean ± SEM; n = 3. **P < 0.01, ***P < 0.001.

Revised version: (A) qRT-PCR analysis validation of E3 ligase knockdown efficiency (shPIRH2, shKCTD11, shMDM2, shCHFR, shUHRF1, shTRIM46) in HeLa cells. Data are mean ± SEM; n = 3. ***P < 0.001. (B) Immunoblotting analysis of HDAC1/2 stability in E3 ligase-knockdown HeLa cells infected with HSV-1 (MOI = 1). (C) Co-IP assay of endogenous HDAC1-MDM2 interaction in HSV-1-infected HeLa cells (MOI = 1; 24 hpi). (D) Co-IP assay of FLAG-HDAC1 with EGFP-MDM2 (WT or ΔRING) in HSV-1-infected HeLa cells (MOI = 1; 24 hpi). (E) Co-IP assay of FLAG-HDAC1 (WT or K74R) with EGFP-MDM2 in HSV-1-infected HeLa cells (MOI = 1; 24 hpi). (F) Ubiquitination assays of FLAG-HDAC1 (WT or K74R) co-expressed with EGFP-MDM2 (WT or ΔRING) in HSV-1-infected HeLa cells (MOI = 1; 24 hpi). (G–I) qRT-PCR analysis of HSV-1 ICP0 (H), ICP8 (I) and gB (J) mRNA levels, normalized to β-actin expression, in Scramble and shMDM2 HeLa cells, and either mock-infected or infected with HSV-1 (MOI = 1) at 24 hpi. Data are mean ± SEM; n = 3. ns > 0.05, *P < 0.05, **P < 0.01, ***P < 0.001. (J) Viral titer analysis in Scramble and shMDM2 HeLa cells, either mock-infected or infected with HSV-1 (MOI = 2) at 24 hpi. t = 0 h defined as time of medium replacement post-adsorption. Data are mean ± SEM; n = 3. **P < 0.01, ***P < 0.001. Note: The numerical values beneath each lane in the Western blot images represent the relative abundance of the corresponding protein bands, as quantified by densitometric analysis and normalized to the signal intensity of the first lane, which serves as the control group.

To characterize this interaction, we focused on the RING finger domain of MDM2, which is essential for its E3 ligase activity. Co-IP analysis demonstrated that HSV-1 infection promoted binding between FLAG-HDAC1 and EGFP-MDM2, but not with a RING-deleted mutant (EGFP-MDM2ΔRING) (Fig. 4D), indicating that the interaction was RING domain-dependent. To identify the region within HDAC1 required for binding, HeLa cells were co-transfected with EGFP-MDM2 and either FLAG-HDAC1 or the ubiquitination-resistant mutant FLAG-HDAC1 K74R. Infection with HSV-1 failed to induce binding between EGFP-MDM2 and the K74R mutant (Fig. 4E), underscoring the essential role of lysine 74 in mediating MDM2-dependent ubiquitination.

Earlier version: Ubiquitination assay further revealed that overexpression of EGFP-MDM2ΔRING did not enhance HSV-1-induced ubiquitination or degradation of HDAC1 (Fig. 4F). Similarly, the FLAG-HDAC1 K74R mutant remained resistant to ubiquitination under all experimental conditions (Fig. 4F). To further validate the role of MDM2, we performed MDM2 knockdown in infected cells; RT–PCR analysis showed that MDM2 knockdown significantly suppressed HSV-1 ICP0, ICP8, and gB mRNA transcription (Fig. 4G–I), and reduced viral titers in plaque assays (Fig. 4J). Collectively, these results demonstrate that HSV-1 induces HDAC1/2 degradation through MDM2-mediated K63-linked ubiquitination.

Revised version: Ubiquitination assay further revealed that overexpression of EGFP-MDM2ΔRING did not enhance HSV-1-induced ubiquitination or degradation of HDAC1 (Fig. 4F). Similarly, the FLAG-HDAC1 K74R mutant remained resistant to ubiquitination under all experimental conditions (Fig. 4F). To further validate the role of MDM2, we performed MDM2 knockdown in infected cells; qRT-PCR analysis showed that MDM2 knockdown significantly suppressed HSV-1 ICP0, ICP8, and gB mRNA transcription (Fig. 4G–I), and reduced viral titers in plaque assays (Fig. 4J). Collectively, these results demonstrate that HSV-1 induces HDAC1/2 degradation through MDM2-mediated K63-linked ubiquitination.

HSV-1 triggers cytoplasmic translocation of HDAC1 to facilitate its degradation

To elucidate the mechanisms by which HSV-1 induces HDAC degradation, we systematically analyzed the subcellular localization of HDAC1. In mock-infected HeLa cells, HDAC1 was predominantly nuclear; however, HSV-1 infection triggered a pronounced translocation of HDAC1 to the cytoplasm (Fig. 5B). This redistribution was completely abolished by treatment with leptomycin B (LMB), a specific inhibitor of CRM1-mediated nuclear export (Fig. 5B). At the same time, CCK-8 assay showed that LMB had no significant effect on cell proliferation (Fig. 5A).

Cytoplasmic translocation enables MDM2-mediated HDAC1 degradation.

Earlier version: (A) HeLa cells were treated with vehicle and LMB (0-60 μM) for 0–48 h. Cell proliferation was analyzed by CCK-8 assay. ns, no significance. (B and C) Immunofluorescence staining of HDAC1 (red) and viral ICP4 (green) in HSV-1-infected HeLa cells (MOI = 1) ± leptomycin B (LMB; 10 nM; 24 h). Scale bar: 10 μm. (C) shows the quantitative analysis of (B), Data are mean ± SEM; n = 100. ***P < 0.001. (D) CHX chase assay measuring HDAC1/2 degradation kinetics in HSV-1-infected HeLa cells (MOI = 1) ± LMB (10 μM). (E) Ubiquitination assay of HDAC1 in HSV-1-infected HeLa cells (MOI = 1) ± LMB (10 μM; 24 h). (F) Immunoblotting analysis of ubiquitination in HDAC1 from cytosolic and nuclear fractions of HeLa cells infected with HSV-1. (G and H) Immunofluorescence analysis of MDM2 (red) and ICP4 (green) in HSV-1-infected HeLa cells (MOI = 1) ± LMB (10 μM; 24 h). Scale bar: 10 μm. (H) shows the quantitative analysis of (G), Data are mean ± SEM; n = 100. ***P < 0.001. (I) Immunoblotting analysis of DDR markers (γ-H2AX) and HDAC1/2 in HSV-1-infected HeLa cells (MOI = 1) ± LMB (10 μM). (J) Viral titers in HSV-1-infected HeLa cells (MOI = 2) ± LMB (10 μM). Data are mean ± SEM; n = 3. *P < 0.05, **P < 0.01, ***P < 0.001 (K) Viral titer analysis in HeLa cells transfected with siHDAC1 transfected with FLAG-HDAC1,FLAG-HDAC1 (K74R) or FLAG-HDAC1 ΔNES, infected with HSV-1 (MOI = 1) at 24 hpi. Data are mean ± SEM; n = 3. **P < 0.01, ***P < 0.001.

Revised version: (A) HeLa cells were treated with vehicle and LMB (0–60 μM) for 0–48 h. Cell proliferation was analyzed by CCK-8 assay. ns, no significance. (B and C) Immunofluorescence staining of HDAC1 (red) and viral ICP4 (green) in HSV-1-infected HeLa cells (MOI = 1) ± leptomycin B (LMB; 10 nM; 24 h). Scale bar: 10 μm. (C) shows the quantitative analysis of (B), Data are mean ± SEM; n = 100. ***P < 0.001. (D) CHX chase assay measuring HDAC1/2 degradation kinetics in HSV-1-infected HeLa cells (MOI = 1) ± LMB (10 μM). (E) Ubiquitination assay of HDAC1 in HSV-1-infected HeLa cells (MOI = 1) ± LMB (10 μM; 24 h). (F) Immunoblotting analysis of ubiquitination in HDAC1 from cytosolic and nuclear fractions of HeLa cells infected with HSV-1. (G and H) Immunofluorescence analysis of MDM2 (red) and ICP4 (green) in HSV-1-infected HeLa cells (MOI = 1) ± LMB (10 μM; 24 h). Scale bar: 10 μm. (H) shows the quantitative analysis of (G), Data are mean ± SEM; n = 100. ***P < 0.001. (I) Immunoblotting analysis of DDR markers (γ-H2AX) and HDAC1/2 in HSV-1-infected HeLa cells (MOI = 1) ± LMB (10 μM). (J) Viral titers in HSV-1-infected HeLa cells (MOI = 2) ± LMB (10 μM). t = 0 h defined as time of medium replacement post-adsorption.Data are mean ± SEM; n = 3. *P < 0.05, **P < 0.01, ***P < 0.001 (K) Viral titer analysis in HeLa cells transfected with siHDAC1 transfected with empty vector, FLAG-HDAC1, FLAG-HDAC1 (K74R) or FLAG-HDAC1 ΔNES, infected with HSV-1 (MOI = 1) at 24 hpi. Data are mean ± SEM; n = 3. **P < 0.01, ***P < 0.001. Note: The numerical values beneath each lane in the Western blot images represent the relative abundance of the corresponding protein bands, as quantified by densitometric analysis and normalized to the signal intensity of the first lane, which serves as the control group.

To determine whether translocation is functionally linked to degradation, we conducted cycloheximide (CHX) chase assays. Degradation of both HDAC1 and HDAC2 began approximately 12 hours after CHX treatment and was markedly accelerated by HSV-1 infection (Fig. 5D). Importantly, LMB treatment restored degradation kinetics to levels comparable to those in mock-infected cells (Fig. 5B). Moreover, under LMB treatment, HSV-1 was unable to induce ubiquitination of endogenous HDAC1 (Fig. 5E). To further validate the subcellular site of HDAC1 ubiquitination, we performed nucleocytoplasmic fractionation followed by immunoblotting; results showed that HSV-1–induced ubiquitination of endogenous HDAC1 occurs specifically in the cytoplasm (Fig. 5F).

We further demonstrated that nuclear export and subsequent proteasomal degradation of HDAC1 are strictly dependent on the E3 ubiquitin ligase MDM2. HSV-1 infection induces CRM1-dependent ubiquitination of endogenous HDAC1, thereby establishing CRM1-mediated nuclear export as an essential prerequisite for MDM2-catalyzed HDAC1 degradation (Fig. 5G, H). Consistent with this, immunoblotting analysis revealed that LMB significantly suppressed HSV-1–induced phosphorylation of p53 and γH2AX (Fig. 5I). Notably, LMB treatment-by inhibiting CRM1-dependent nuclear export and thereby preventing HDAC1 degradation–markedly impaired HSV-1 replication (Fig. 5F). To rigorously test whether HDAC1 nuclear export is the critical, rate-limiting step enabling its ubiquitination and degradation, we employed NESmapper to identify a canonical nuclear export signal (NES) in HDAC1 (residues 207–216). We then generated an HDAC1 ΔNES deletion mutant and performed functional rescue experiments in HDAC1-knockdown cells by reconstituting expression of HDAC1 WT, the ubiquitination-deficient mutant HDAC1 K74R, or FLAG-tagged HDAC1 ΔNES. Only HDAC1 WT fully restored the pro-viral phenotype associated with HDAC1 depletion; in contrast, neither HDAC1 K74R nor HDAC1 ΔNES rescued viral replication, and both exerted potent antiviral activity–significantly suppressing HSV-1 proliferation relative to controls (Fig. 5K). Collectively, these findings establish that HSV-1 co-opts the CRM1 nuclear export machinery to translocate HDAC1 to the cytoplasm, where it becomes accessible to MDM2 for ubiquitination and subsequent proteasomal degradation, thereby modulating host antiviral defense.

Discussion

HSV-1 manipulates host cellular pathways to facilitate its replication, latency, and reactivation [26]. In this study, we demonstrate that alphaherpesvirus targets class I histone deacetylases HDAC1 and HDAC2 for MDM2-mediated K63-linked polyubiquitination and subsequent proteasomal degradation. This degradation induces histone hyperacetylation, chromatin relaxation, and enhanced viral replication–revealing a mechanism through which alphaherpesvirus reprograms the host epigenetic landscape to optimize viral replication (Fig. 6).

A schematic model showing the targeted cytosolic degradation of class I histone deacetylases is essential for efficient alphaherpesvirus replication.

In contrast, HSV-1 actively targets HDAC1 and HDAC2 for proteasomal degradation to facilitate lytic replication–a strategy that stands in marked contrast to human cytomegalovirus (HCMV), which relies on HDAC activity to maintain latency [27]. This divergence underscores a fundamental difference in epigenetic regulation among herpesviruses. Our data demonstrate that HSV-1–mediated HDAC1/2 degradation initiates as early as 4–6 hours post-infection, coinciding with the onset of immediate-early gene expression and preceding peak viral DNA synthesis. These findings support a context-dependent, dual functional role for HDAC1/2: their enzymatic activity helps enforce transcriptional silencing during latency, whereas their timely removal is essential for chromatin remodeling and efficient lytic gene expression.

We identified MDM2 [28] as the principal E3 ubiquitin ligase responsible for ubiquitinating HDAC1 at lysine 74 and HDAC2 at lysine 75, leading to their proteasomal degradation. This mechanism is distinct from other E3 ligases, including PIRH2 [29] and TRIM46 [30], and critically depends on the RING domain of MDM2 [31]. While K63-linked ubiquitination is typically associated with non-proteolytic functions, our work establishes its novel role in mediating HDAC1/2 degradation–uncovering a previously unrecognized viral pathogenesis pathway. Targeting this axis using MDM2 inhibitors or inhibiting nuclear export of HDAC1/2 significantly restricts viral replication, suggesting that host-directed antiviral strategies may help overcome resistance to direct-acting agents [31].

Depletion of HDAC1/2 increases acetylation at histone residues associated with open chromatin and active transcription. Similar chromatin remodeling has been observed in other viral infections [32], underscoring its conserved role in viral gene regulation. Our results are consistent across multiple cell types, including HeLa, 3D4/21, and murine liver, emphasizing the central importance of HDAC1/2 suppression in HSV-1 biology. Notably, HDAC1/2 degradation activates DDR, as indicated by increased γ-H2AX foci and phosphorylation of ATM, ATR, Chk1, and Chk2. This aligns with previous reports that HSV-1 exploits DDR signaling to enhance replication [33-35]. Importantly, HDAC1/2 undergo CRM1-dependent nuclear export prior to degradation–a process inhibited by leptomycin B–highlighting the essential role of subcellular localization in viral manipulation of host proteins.

HSV-1–mediated HDAC1/2 degradation likely supports viral replication through multiple mechanisms, including evasion of nuclear antiviral sensors and stimulation of viral gene expression. Targeting the MDM2–HDAC1/2 axis represents a promising antiviral strategy. MDM2 inhibitors reduce HSV-1 titers [36], suggesting that modulating ubiquitination pathways could complement existing therapies. Beyond herpesviruses, the MDM2–HDAC axis may be exploited by other DNA viruses that manipulate chromatin and DDR [32], indicating a possible common host vulnerability. However, broad HDAC inhibition may unintentionally enhance viral replication, underscoring the need for targeted therapeutic approaches [37].

This study primarily employed in vitro models (HeLa cells), which lack the neuronal environment essential for studying HSV-1 latency. Future work should validate these findings in primary neurons and trigeminal ganglia [38]. Additional studies are also needed to determine whether HDAC1/2 degradation occurs during reactivation or is limited to lytic infection [7], and to explore the interplay between HDAC degradation, viral tegument proteins, and host stress responses. In vivo evaluation of inhibitors targeting the MDM2–HDAC1/2 pathway will be crucial for assessing their therapeutic potential.